[1733939346] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 401 Unauthorized - 695 bytes in 0.001 second response time
[1733939406] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 401 Unauthorized - 695 bytes in 0.001 second response time
[1733939466] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 401 Unauthorized - 695 bytes in 0.001 second response time
[1733939526] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 401 Unauthorized - 695 bytes in 0.001 second response time
[1733948102] SERVICE ALERT: dbserver;MYSQLD;CRITICAL;SOFT;1;(No output on stdout) stderr: execvp(/opt/nagios/libexec/check_mysqld.pl, ...) failed. errno is 2: No such file or directory
[1733948222] SERVICE ALERT: dbserver;MYSQLD;CRITICAL;SOFT;2;(No output on stdout) stderr: execvp(/opt/nagios/libexec/check_mysqld.pl, ...) failed. errno is 2: No such file or directory
[1733948342] SERVICE ALERT: dbserver;MYSQLD;CRITICAL;HARD;3;(No output on stdout) stderr: execvp(/opt/nagios/libexec/check_mysqld.pl, ...) failed. errno is 2: No such file or directory
[1733952258] SERVICE ALERT: dbserver;MYSQLD;OK;HARD;1;Uptime: 2282  Threads: 2  Questions: 11  Slow queries: 0  Opens: 138  Flush tables: 3  Open tables: 57  Queries per second avg: 0.004
[1733984601] SERVICE ALERT: webserver;PING;CRITICAL;SOFT;1;(Service check timed out after 32171.66 seconds)
[1733984665] SERVICE ALERT: webserver;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.07 ms
[1733985148] SERVICE ALERT: dbserver;MYSQLD;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 172.19.0.4 port 5666: Connection refused
[1733985268] SERVICE ALERT: dbserver;MYSQLD;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 172.19.0.4 port 5666: Connection refused
[1733985388] SERVICE ALERT: dbserver;MYSQLD;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 172.19.0.4 port 5666: Connection refused
[1733987788] SERVICE ALERT: dbserver;MYSQLD;OK;HARD;1;Uptime: 2672  Threads: 2  Questions: 2  Slow queries: 0  Opens: 119  Flush tables: 3  Open tables: 38  Queries per second avg: 0.000
[1733993618] SERVICE ALERT: dbserver;MYSQLD;OK;HARD;1;Uptime: 285  Threads: 2  Questions: 5  Slow queries: 0  Opens: 136  Flush tables: 3  Open tables: 55  Queries per second avg: 0.017
[1733993918] SERVICE ALERT: dbserver;MYSQLD;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 172.20.0.3 port 5666: Connection refused
[1733994038] SERVICE ALERT: dbserver;MYSQLD;CRITICAL;SOFT;2;(No output on stdout) stderr: connect to address 172.20.0.3 port 5666: Connection refused
[1733994158] SERVICE ALERT: dbserver;MYSQLD;CRITICAL;HARD;3;(No output on stdout) stderr: connect to address 172.20.0.3 port 5666: Connection refused